DNA tensiometer reveals catch-bond detachment kinetics of kinesin-1, -2 and -3

Author response:

The following is the authors’ response to the current reviews.

Reviewer #1 (Public review):

I am not fully convinced about the responses from authors, so I would like to retain my original assessment of the paper. The same may be made available for public viewing, along with the responses of the authors. Readers can go through both and form their opinion.

Unfortunately, this response from Reviewer 1 impacted the Assessment Statement but did not provide specific points for us to address. In the first round, the concerns of Reviewer 1 were: 1) the validity of the WLC prediction; 2) the claim that catch-bond measurements are generally made with superstall loads; 3) the role of vertical forces for dynein and a question about the orientation of the forces for kinesin; and 4) a request that we repeat the study using dynein. In rereading our responses to points 2-4 following our first revision, we felt that there were no unresolved issues around those points that affect our conclusions in any way. However, for point 1 regarding the validity of the WLC prediction, we had responded only in the reviewer response letter, and both reviewer 2 and the editors felt that there were points that we had addressed in the response letter that should be incorporated into the revised manuscript. Therefore, to clarify Reviewer 1’s question, we revised the text to address why we were justified to approximate the dsDNA force-extension curve using a WLC model with a 50 nm persistence length and why the precise shape of the force-extension curve has no impact on our conclusions.

Reviewer #2 (Public review):

The authors extensively entered into a scientific debate with the reviewers in their Response Letter. This led to a few changes and some (limited) new data in the manuscript. This is great and did improve the manuscript.

However, in the view of this reviewer, (i) a significant number of responses fall short of actually addressing the concerns of the three reviewers (e.g. wrt using the same kinesin-1 neck-coil domains for all motors) and or (ii) a significant number of arguments now only occur in the response letter but not in the manuscript. The authors may check themselves critically for both. In principle, each longer discussion in the response letter warrants mentioning the appropriate facts and arguments in the main text of the manuscript.

Based on this feedback, the first change we made was to rewrite the section justifying our choice of using a common coiled-coil dimerization domain for the three motors. Secondly, we went through our responses to all three reviewers to identify any instances where we either didn’t fully address the reviewer concerns or we provided arguments in the response letter but did not add corresponding text in the manuscript.

Reviewer #3 (Public review):

The authors attribute the differences in the behaviour of kinesins when pulling against a DNA tether compared to an optical trap to the differences in the perpendicular forces. However, the compliance is also much different in these two experiments. The optical trap acts like a ~ linear spring with stiffness ~ 0.05 pN/nm. The dsDNA tether is an entropic spring, with negligible stiffness at low extensions and very high compliance once the tether is extended to its contour length (Fig. 1B). The effect of the compliance on the results is not fully considered in the manuscript.

In our first revision we added a paragraph in the ‘Geometry Calculations section of the Supplementary Methods addressing the dsDNA stiffness and comparing it to an optical trap. We considered moving this paragraph to the main text but decided against it because we felt it interrupted the flow of the Discussion. Instead, we expanded and clarified this paragraph to more specifically address the stiffness question. The paragraph with revised text now reads as follows:

“Another consideration when comparing the DNA tensiometer to optical trap measurements is the relative stiffness of the trap and dsDNA. Optical traps stiffnesses are generally in the range of 0.05 pN/nm [13,14]. To calculate the predicted stiffness of the dsDNA spring, we computed the slope of theoretical force-extension curve in Fig. 1B. The stiffness is highly nonlinear and is <0.001 pN/nM below 650 nm extension. We compare motor performance under this low stiffness regime to the unloaded case in Fig. 3. In contrast, at the predicted stall force of 6 pN (960 nm extension), the dsDNA stiffness is ~0.2 pN/nm, which is stiffer than most optical traps, but it is similar to the estimated 0.3 pN/nm stiffness of kinesin motors themselves [13,14]. An 8 nm step at the 0.2 pN/nm stiffness of the dsDNA leads to a 1.6 pN jump in force and at the 0.05 pN/nm stiffness of an optical trap leads to a 0.4 pN jump in force; this is important because it means that in both cases the motors are likely dynamically stepping at stall. Because both experimental approaches allow for dynamic stepping at stall and because the stiffnesses of the instrument in both cases are less than the motor stiffness, there is no reason to expect that differences in stiffness between optical traps and the dsDNA spring lead to different motor detachment kinetics.”

In the main text, we now address this compliance point in the ‘Comparison to previous work’ section of the Discussion:

“stiffness differences are an unlikely explanation because at stall the stiffness of the DNA tether (~4 fold stiffer than optical tweezer) is still sufficiently low to allow for dynamic motor stepping at stall, and in any case it is still below the estimated motor stiffness (see Geometry Calculations in Supplementary methods).”.

There were two points the reviewer felt we had sufficiently addressed. They were presented in the second review as a reiteration of the first review comments with a sentence appended, and are reproduced here. We added no new text based on these two points:

In the single-molecule extension traces (Fig. 1F-H; S3), the kinesin-2 traces often show jumps in position at the beginning of runs (e.g. the four runs from ~4-13 s in Fig. 1G). These jumps are not apparent in the kinesin-1 and -3 traces. What is the explanation? Is kinesin-2 binding accelerated by resisting loads more strongly than kinesin-1 and -3? In their response, the authors provide an explanation of the appearance of jumps due to limited imaging speeds. The authors state that the qualitative difference in the kinesin-2 traces compared to the kinesin-1 and -3 traces may be due to the specific rebinding kinetics of kinesin-2.

When comparing the durations of unloaded and stall events (Fig. 2), there is a potential for bias in the measurement, where very long unloaded runs cannot be observed due to the limited length of the microtubule (Thompson, Hoeprich, and Berger, 2013), while the duration of tethered runs is only limited by photobleaching. Was the possible censoring of the results addressed in the analysis? The authors addressed this concern by applying a Markov model to estimate the duration parameter.

There was one final point from Reviewer 3 in the first round of reviews that we had addressed in the reviewer response (and that the reviewer was satisfied with), but we did not incorporate into the manuscript. Based on the suggestion from Reviewer 2 and the editors that we incorporate more from our responses to reviewers into the manuscript, we added new text on this point. That point (with the new sentence in the second review underlined), our response from first revision, and our response for this second revision are given below:

The mathematical model is helpful in interpreting the data. To assess how the "slip" state contributes to the association kinetics, it would be helpful to compare the proposed model with a similar model with no slip state. Could the slips be explained by fast reattachments from the detached state? In their response, the authors addressed this question by explaining that a three-state model is required to model the recovery time distributions.

In the model, the slip state and the detached states are conceptually similar; they only differ in the sequence (slip to detached) and the transition rates into and out of them. The simple answer is: yes, the slips could be explained by fast reattachments from the detached state. In that case, the slip state and recovery could be called a “detached state with fast reattachment kinetics”. However, the key data for defining the kinetics of the slip and detached states is the distribution of Recovery times shown in Fig. 4D-F, which required a triple exponential to account for all of the data. If we simplified the model by eliminating the slip state and incorporating fast reattachment from a single detached state, then the distribution of Recovery times would be a single-exponential with a time constant equivalent to t1, which would be a poor fit to the experimental distributions in Fig. 4D-F.

Reviewer 3 noted that they were satisfied with our explanation of this point. However, based on Reviewer 2’s suggestion that we incorporate more of our responses into the text of the manuscript, we added the following clarification point in the model section of the Results:

“We note that recapitulating the tri-exponential restart time distribution in Figure 4D-F required this slip/detached formulation and that lumping all events into a single detached state resulted in a single-exponential distribution of recovery times.”

Changes from the earlier version to the revised version

Abstract

Earlier version: Bidirectional cargo transport by kinesin and dynein is essential for cell viability and defects are linked to neurodegenerative diseases. Computational modeling suggests that the load-dependent off-rate is the strongest determinant of which motor ‘wins’ a kinesin-dynein tug-of-war, and optical tweezer experiments find family-dependent differences in the sensitivity of detachment to load, with kinesin-3 > kinesin-2 > kinesin-1. However, in reconstituted kinesin-dynein pairs vitro, all three kinesin families compete nearly equally well against dynein. Modeling and experiments have confirmed that vertical forces inherent to the large trapping beads enhance kinesin-1 dissociation rates. In vivo, vertical forces are expected to range from negligible to dominant, depending on cargo and microtubule geometries. To investigate the detachment and reattachment kinetics of kinesin-1, 2 and 3 motors against loads oriented parallel to the microtubule, we created a DNA tensiometer comprising a DNA entropic spring attached to the microtubule on one end and a motor on the other. Kinesin dissociation rates at stall were slower than detachment rates during unloaded runs, and the complex reattachment kinetics were consistent with a weakly-bound ‘slip’ state preceding detachment. Kinesin-3 behaviors under load suggested that long KIF1A run lengths result from the concatenation of multiple short runs connected by diffusive episodes. Stochastic simulations were able to recapitulate the load-dependent detachment and reattachment kinetics for all three motors and provide direct comparison of key transition rates between families. These results provide insight into how kinesin-1, -2 and -3 families transport cargo in complex cellular geometries and compete against dynein during bidirectional transport.

Revised version: Bidirectional cargo transport by kinesin and dynein is essential for cell viability, and defects are linked to neurodegenerative disease. Computational models predict that load-dependent motor detachment strongly determines the outcome of kinesin–dynein tug-of-war, with kinesin-3 and kinesin-2 more load-sensitive than kinesin-1. Yet reconstituted assays show that all three kinesin families compete similarly well against dynein. Previous work demonstrated that vertical forces from optical trapping assays can enhance kinesin-1 dissociation, suggesting that motor behavior may depend strongly on cargo geometry. To measure kinesin detachment and reattachment kinetics under forces applied parallel to the microtubule, we developed a DNA-based tensiometer using an entropic DNA spring linking motors to microtubules. For kinesin-1, -2, and -3, dissociation rates at stall were slower than during unloaded motion, and reattachment kinetics were consistent with a weakly bound slip state preceding detachment. Kinesin-3 behavior further suggested that long KIF1A run lengths arise from multiple short runs connected by diffusive episodes. Stochastic simulations reproduced the measured load-dependent kinetics and enabled direct comparison of transition rates among kinesin families. These results provide insight into how kinesin-1, -2 and -3 transport cargo in complex cellular geometries and compete against dynein during bidirectional transport.

Introduction

Bidirectional cargo transport by kinesin and dynein motors is essential for cell viability, and disruptions in transport are linked to neurological diseases including hereditary spastic paraplegia, microcephaly and amyotrophic lateral sclerosis 1–8. It has been established that kinesin and dynein, which move in opposite directions along microtubules, are often bound simultaneously to the same cargo 9–12. This has led to the ‘tug-of-war’ model, in which the direction of cargo movement is determined by which team of motors dominates 9,13–17. How well motors compete is determined by their load-dependent motor properties along with multiple regulation mechanisms, many of which are still emerging 13,18–21. Furthermore, the large range of cargo sizes and the complexity of microtubule organization in cells means that motors are subjected to forces both parallel and perpendicular to their microtubule track, which can have differing effects on their mechanochemistry.

Intuitively, a motor’s effectiveness in transporting cargo rests on its ability to remain bound to its microtubule track. Consistent with this, computational simulations have found that the load-dependent off-rate of a motor is the most important determinant of how well a kinesin competes against dynein in bidirectional transport 22,23. Single-bead optical tweezers have found that the transport motors kinesin-1, -2, and -3 all act as slip bonds, defined as load accelerating their detachment rate. Their propensity to detach under load varies strongly by family, with relative load sensitivity kinesin-3 > kinesin-2 > kinesin-1 24–27. Based on this behavior, it was surprising that when kinesin-1 was linked to dynein, complexes moved at near-zero speeds for up to tens of seconds, much longer than predicted based on previously measured kinesin-1 off-rates 25,28,29. Moreover, kinesin-1, -2, and -3 all fared equally well against dynein, contrary to the differing load-dependent detachment rates measured in single-bead optical tweezer experiments 30.

Earlier version: Recent work suggests a solution to this paradox, namely that the ∼micron scale beads used for optical trapping result in significant forces oriented perpendicular to the microtubule as the motor pulls against the force of the trap. First, the load-dependent dissociation rate from single-bead optical trapping was accounted for by a model in which the effects of horizontal loads on detachment is highly asymmetric and vertical loads play a dominant role in detachment, particularly against hindering loads 31. Second, when a three-bead optical trapping geometry was used (in which the motor is raised up on a pedestal bead and the microtubule was held by beads attached to either end of the microtubule) motors remained bound longer than in the single-bead geometry 24,32. Third, when kinesin-1 motors were connected to a microtubule by a micron-long segment of DNA, very long residence times were observed, consistent with catch-bond behavior, defined as the off-rate slowing with load 33. Fourth, when kinesin-1 was connected to a bead through a micron-long segment of DNA and hydrodynamic forces were imposed on the bead, motor interaction times were insensitive to hindering loads up to 3 pN, indicative of an ideal-bond 34. In cells, kinesin and dynein transport cargoes that range from tens of nm in diameter (like vesicles), where motor forces are expected to be aligned parallel to the microtubule, up to several microns (like mitochondria and nuclei), where vertical forces are expected to be much larger. Thus, understanding the influence of vertical and horizontal forces on transport motors is important for understanding the mechanics underlying bidirectional transport in cells.

Revised version: Recent work suggests a solution to this paradox, namely that the ∼micron scale beads used for optical trapping result in significant forces oriented perpendicular to the microtubule as the motor pulls against the force of the trap. First, the load-dependent dissociation rate from single-bead optical trapping was accounted for by a model in which the effects of horizontal loads on detachment is highly asymmetric and vertical loads play a dominant role in detachment, particularly against hindering loads 31.

Added in the revised version: Second, when a three-bead optical trapping geometry was used (in which the motor is raised up on a pedestal bead and the microtubule was held by beads attached to either end of the microtubule) motors remained bound longer than in the single-bead geometry 24,32. Third, when kinesin-1 motors were connected to a microtubule by a micron-long segment of DNA, very long residence times were observed, consistent with catch-bond behavior, defined as the off-rate slowing with load 33. Fourth, when kinesin-1 was connected to a bead through a micron-long segment of DNA and hydrodynamic forces were imposed on the bead, motor interaction times were insensitive to hindering loads up to 3 pN, indicative of an ideal-bond 34. In cells, kinesin and dynein transport cargoes that range from tens of nm in diameter (like vesicles), where motor forces are expected to be aligned parallel to the microtubule, up to several microns (like mitochondria and nuclei), where vertical forces are expected to be much larger. Thus, understanding the influence of vertical and horizontal forces on transport motors is important for understanding the mechanics underlying bidirectional transport in cells.

The goal of the present study was to characterize the load-dependent detachment kinetics of kinesin-1, -2 and -3 motors in a geometry that eliminates vertical forces inherent in traditional optical trapping studies. Building on previous approaches, we used double stranded DNA (dsDNA), which acts as an entropic spring to resist the pN-level forces generated by the motors 33–38. We found that kinesin-1, -2, and -3 all remained at stall for multiple seconds before releasing, durations that are substantially longer than the unloaded run times for kinesin-1 and -2. This behavior of slower off-rates under load is defined as a ‘catch bond’, and contrasts with the normal ‘slip bond’ behavior load-accelerated off-rates generally observed for kinesin 25,39,40. Following the termination of a stall, motors reengaged with the microtubule with complex kinetics that were consistent with a ‘slip’ state that preceded full detachment. To compare the key transition rates that determined the family-specific motor behaviors, we developed a stochastic model that was able to recapitulate the experimental results for all three motors.

Results

Constructing a motor-bound DNA tensiometer

Earlier version: To study motor performance against a resistive load oriented parallel to the microtubule, we constructed a DNA tensiometer consisting of a ∼ 1 μm strand of dsDNA attached to the microtubule on one end and a motor on the other (Figure 1A). We used TIRF microscopy to visualize the motor moving against the entropic elasticity of the DNA spring. Due to the nonlinear elasticity of the DNA (Figure 1B) 35,36, the motor moves under minimal load until it stretches the DNA to near its contour length, at which point it stalls (Figure 1C-E).

Revised version: To study motor performance against a resistive load oriented parallel to the microtubule, we constructed a DNA tensiometer consisting of a ∼ 1 μm strand of dsDNA attached to the microtubule on one end and a motor on the other (Figure 1A). We used TIRF microscopy to visualize the motor moving against the entropic elasticity of the DNA spring. Experimental and theoretical studies have established that the force-extension profile of dsDNA is well fit by a Worm-Like Chain model with a persistence length of ∼50 nm 36,41–43. As expected from this nonlinear elasticity (Figure 1B), we observed motors moving at near their unloaded velocity to near the ∼ 1 μm contour length of the DNA, at which point they stalled and eventually detached (Figure 1C-E).

Experimental Design and Raw Data from Motor-DNA Tensiometers.

Earlier version: (A) Schematic of a motor-DNA tensiometer, consisting of a dsDNA (burgundy) connected on one end to a kinesin motor through a complimentary oligo (blue), and on the other end to the MT using biotin-avidin (tan and gray, respectively). A Qdot functionalized with GFP binding protein nanobodies is attached to the motor’s GFP tag and used to track motor position. (Not to scale; motor and Qdot are both ∼20 nm and DNA is ∼1 micron) (B) Predicted force extension curve for a worm-like chain 3009 bp dsDNA based on a 50 nm persistence length. (C) Representative kymographs of motor-DNA tensiometers for kinesins-1, -2 and -3. (D) Enlarged kymograph showing diffusion around the origin, ramp, and stall. (E) Example distance vs. time trace (kinesin-3), highlighting detached durations (red), ramps and stalls (black) where the motor has pulled the DNA taut, and transient slips during stall (red). (F-H) Representative distance vs. time plots for kinesin-1 (F), kinesin-2 (G) and kinesin-3 (H), corresponding to the kymographs in (C). Further examples are shown in Figure S3.

Revised version: (A) Schematic of a motor-DNA tensiometer, consisting of a dsDNA (burgundy) connected on one end to a kinesin motor through a complimentary oligo (blue), and on the other end to the MT using biotin-avidin (tan and gray, respectively). A Qdot functionalized with GFP binding protein nanobodies is attached to the motor’s GFP tag and used to track motor position. (Not to scale; motor and Qdot are both ∼20 nm and DNA is ∼1 micron) (B) Predicted force extension curve for a worm-like chain 3009 bp dsDNA based on a 50 nm persistence length. We note that our analysis of family-dependent mechanochemistry does not depend on the precise shape of the force-extension curve, only that the motors reach stall. (C) Representative kymographs of motor-DNA tensiometers for kinesins-1, -2 and -3. (D) Enlarged kymograph showing diffusion around the origin, ramp, and stall. (E) Example distance vs. time trace (kinesin-3), highlighting detached durations (red), ramps and stalls (black) where the motor has pulled the DNA taut, and transient slips during stall (red). (F-H) Representative distance vs. time plots for kinesin-1 (F), kinesin-2 (G) and kinesin-3 (H), corresponding to the kymographs in (C). Further examples are shown in Figure 1 – figure supplement 3.

Earlier version: Our DNA-motor tensiometer consists of a 3,009 bp (999 nm contour length) dsDNA ‘spring’ that was synthesized by PCR using a biotinylated forward primer for attachment to the microtubule and a reverse primer containing a 3’ overhang for motor attachment (details in Methods). We investigated members of the three dominant families of kinesin transport motors, kinesin-1 (Drosophila melanogaster KHC), kinesin-2 (Mus musculus Kif3A), and kinesin-3 (Rattus norvegicus Kif1A) used in our previous kinesin-dynein study 30. In each case, the motor and neck linker domains were fused to the stable neck-coil domain (residues 345-406) of kinesin-1, followed by EGFP, a SNAP tag, and His6 tag, as described previously 30. This dimerization strategy avoids any autoinhibition and family-dependent differences in neck-coil stability, thus enabling the most direct comparison of family-dependent motor properties. Motors were conjugated to an oligonucleotide complimentary to the 3’ overhang of the dsDNA spring via their C-terminal SNAP tag. The DNA tensiometer complex (Figure 1A) was created in a flow cell by sequentially flowing in biotinylated microtubules, neutravidin, biotinylated dsDNA, and Qdot-functionalized motors containing the complimentary oligo (described fully in Methods).

Revised version: Our DNA-motor tensiometer consists of a 3,009 bp (999 nm contour length) dsDNA ‘spring’ that was synthesized by PCR using a biotinylated forward primer for attachment to the microtubule and a reverse primer containing a 3’ overhang for motor attachment (details in Methods). We investigated members of the three dominant families of kinesin transport motors, kinesin-1 (Drosophila melanogaster KHC), kinesin-2 (Mus musculus Kif3A), and kinesin-3 (Rattus norvegicus Kif1A). The coiled-coil stability and degree of autoinhibition differ between families and can strongly affect motor function 44–49. Because our goal was to compare the mechanochemical properties of the motor domains, the head and neck linker of each motor was fused to the stable neck-coil domain (residues 345-406) of kinesin-1, followed by EGFP, a SNAP tag, and His6 tag. All of these constructs have been previously characterized by single-molecule TIRF, ATPase, stopped-flow, and optical tweezers studies, as well as by connecting them to a dynein-dynactin-BicD2 complex 24,30,50–53. This dimerization approach enables direct comparison to that body of published work and it eliminates the possibility that family-specific differences result from differences in the stability of the neck-coil domain under load or differences in the degree of autoinhibition. Motors were conjugated to an oligonucleotide complimentary to the 3’ overhang of the dsDNA spring via their C-terminal SNAP tag. The DNA tensiometer complex (Figure 1A) was created in a flow cell by sequentially flowing in biotinylated microtubules, neutravidin, biotinylated dsDNA, and Qdot-functionalized motors containing the complimentary oligo (described fully in Methods).

Earlier version: The resulting dsDNA tensiometer kymographs (Figure 1C) show a reproducible behavior of moving, stalling, and returning to origin multiple times, which contrasts with the singular attachment, unidirectional movement and detachment of motors not bound by DNA (Figure S1). Because motors are tethered to the microtubule by the flexible DNA, large fluctuations around the origin are observed when the motor is detached (Figure 1C-H). Consistent with these fluctuations, initial attachment points were variable and roughly normally distributed with a standard deviation of 145 nm (Figure S2). Upon engagement with the microtubule, the motor walks at a steady velocity, consistent with the expected nonlinear stiffness of the dsDNA tether (Figure 1B), until it either disengages or reaches a stall state. Stalls are terminated either by the motor slipping backwards a short distance and restarting a new ramp, or by the motor fully disengaging and returning to the baseline (Figure 1E-H and Figure S3). To confirm that motors are indeed extending the DNA and that Qdots are not enabling multi-motor assemblies, we incorporated Cy5-dCTP into the dsDNA and left the Qdots out of the reaction. In this case, clear extensions of the DNA spring could be observed, and the stall durations were of similar duration (Figure S4). In all subsequent experiments dsDNA was labeled with a low concentration of Cy5-dCTP to confirm colocalization of the DNA and the microtubule before collecting tensiometer data.

Revised version: The resulting dsDNA tensiometer kymographs (Figure 1C) show a reproducible behavior of moving, stalling, and returning to origin multiple times, which contrasts with the singular attachment, unidirectional movement and detachment of motors not bound by DNA (Figure 1 – figure supplement 1). Because motors are tethered to the microtubule by the flexible DNA, large fluctuations around the origin are observed when the motor is detached (Figure 1C-H). Consistent with these fluctuations, initial attachment points were variable and roughly normally distributed with a standard deviation of 145 nm (Figure 1 – figure supplement 2). Upon engagement with the microtubule, the motor walks at a steady velocity, consistent with the expected nonlinear stiffness of the dsDNA tether (Figure 1B), until it either disengages or reaches a stall state. Stalls are terminated either by the motor slipping backwards a short distance and restarting a new ramp, or by the motor fully disengaging and returning to the baseline (Figure 1E-H and Figure 1 – figure supplement 3). To confirm that motors are indeed extending the DNA and that Qdots are not enabling multi-motor assemblies, we incorporated Cy5-dCTP into the dsDNA and left the Qdots out of the reaction. In this case, clear extensions of the DNA spring could be observed, and the stall durations were of similar duration (Figure 1 – figure supplement 4). In all subsequent experiments dsDNA was labeled with a low concentration of Cy5-dCTP to confirm colocalization of the DNA and the microtubule before collecting tensiometer data.

Kinesin-1 and -2 act as catch-bonds at stall

Earlier version: The first question we addressed was: what are the detachment rates of kinesin-1, -2 and -3 motors at stall? The load-dependence of protein-protein interactions can be described as a slip-bond 41, defined as a faster off-rate under load; an ideal bond, defined as an off-rate that is independent of load; or a catch-bond, in which the off-rate is slower under load 39. Single-bead optical trapping studies consistently find slip-bond characteristics for kinesin-1, 2 and 3 25,27,42, whereas dynein off-rates have been described as a slip-bond or catch-bond 43-46.

Revised version: The first question we addressed was: what are the detachment rates of kinesin-1, -2 and -3 motors at stall? The load-dependence of protein-protein interactions can be described as a slip-bond 54, defined as a faster off-rate under load; an ideal bond, defined as an off-rate that is independent of load; or a catch-bond, in which the off-rate is slower under load 39. Single-bead optical trapping studies consistently find slip-bond characteristics for kinesin-1, 2 and 3 25,27,55, whereas dynein off-rates have been described as a slip-bond or catch-bond 56–60.

Earlier version: We define stall duration as the time that a motor stalls against the hindering load of fully extended DNA, without further detectable stepping. Stalls are terminated by the motor detectably (>60 nm) slipping backwards or by disengaging and returning to the origin (Figure 1E). Although we don’t directly measure the stall force, based on the predicted force-extension curve of the dsDNA (Figure 1B), the displacements are consistent with the 4-6 pN stall forces for kinesin-1, -2 and -3 measured using optical traps 27,47-51. Stall durations were compared to the unloaded single-motor run durations determined from TIRF kymograph analysis (Figure S1).

Revised version: We define stall duration as the time that a motor stalls against the hindering load of fully extended DNA, without further detectable stepping. Stalls are terminated by the motor detectably (>60 nm) slipping backwards or by disengaging and returning to the origin (Figure 1E). Although we don’t directly measure the stall force, based on the predicted force-extension curve of the dsDNA (Figure 1B), the displacements are consistent with the 4-6 pN stall forces for kinesin-1, -2 and -3 measured using optical traps 27,61–65. Stall durations were compared to the unloaded single-motor run durations determined from TIRF kymograph analysis (Figure 1 – figure supplement 1).

Earlier version: To compare unloaded off-rates to off-rates at stall, cumulative distributions of the run and stall durations were plotted for each motor and fit with a single exponential function (Figure 2). The kinesin-1 tensiometer stall duration time constant was 3.01 s, with 95% confidence intervals (CI) of 2.30 to 3.79 s determined via bootstrapping in MEMLET 52 with 1000 iterations (N= 78 stalls). In contrast, the kinesin-1 unloaded run duration time constant, measured by a traditional TIRF assay, was 1.04 s, (Figure 2A; Table S1). Stall durations longer than unloaded run durations indicate that load slows the off-rate, the definition of a catch-bond 39. Similarly, the kinesin-2 tensiometer stall duration time constant of 2.83 s was longer than its unloaded run duration of 1.07 s, also indicating a catch-bond (Figure 2B). Conversely, the kinesin-3 tensiometer stall duration time constant of 1.89 s was shorter than its unloaded run duration of 2.74 s, indicating a slip-bond characteristic by this definition.

Revised version: To compare unloaded off-rates to off-rates at stall, cumulative distributions of the run and stall durations were plotted for each motor and fit with a single exponential function (Figure 2). The kinesin-1 tensiometer stall duration time constant was 3.01 s, with 95% confidence intervals (CI) of 2.30 to 3.79 s determined via bootstrapping in MEMLET 66 with 1000 iterations (N= 78 stalls). In contrast, the kinesin-1 unloaded run duration time constant, measured by a traditional TIRF assay, was 1.04 s, (Figure 2A; Table S1). Stall durations longer than unloaded run durations indicate that load slows the off-rate, the definition of a catch-bond 39. Similarly, the kinesin-2 tensiometer stall duration time constant of 2.83 s was longer than its unloaded run duration of 1.07 s, also indicating a catch-bond (Figure 2B). Conversely, the kinesin-3 tensiometer stall duration time constant of 1.89 s was shorter than its unloaded run duration of 2.74 s, indicating a slip-bond characteristic by this definition.

Tensiometer Stall Durations Indicate Catch-bond Behavior for Kinesin-1 and -2.

Earlier version: Tensiometer stall durations are plotted for (A) kinesin-1 (blue), (B) kinesin-2 (purple), and (C) kinesin-3 (green). Unloaded run durations for each motor are plotted in gray. Distributions were fit with a single exponential function using MEMLET to generate time constants, representing the mean durations. (D) Comparison of unloaded and stall durations for the three motors, with error bars indicating 95% CI. Stall durations >20s were excluded from the fit (three events for kinesin-1 and two events for kinesin-2). All fit parameters are given in Table S1 and bi-exponential fits of all data including >20 s are shown in Figure S5.

Revised version: Tensiometer stall durations are plotted for (A) kinesin-1 (blue), (B) kinesin-2 (purple), and (C) kinesin-3 (green). Unloaded run durations for each motor are plotted in gray. Distributions were fit with a single exponential function using MEMLET to generate time constants, representing the mean durations. (D) Comparison of unloaded and stall durations for the three motors, with error bars indicating 95% CI. Stall durations >20s were excluded from the fit (three events for kinesin-1 and two events for kinesin-2). All fit parameters are given in Table S1 and bi-exponential fits of all data including >20 s are shown in Figure 2 – figure supplement 1.

Earlier version: We carried out two additional control experiments. First, to confirm that the neutravidin used to link the DNA to the microtubule wasn’t affecting kinesin motility, we analyzed the run durations of kinesin-1 motors on neutravidin-coated microtubules and found no significant change in the unloaded run duration (Fig. S6). Second, we measured the run duration of kinesin-1 linked to a DNA tether that was not bound to the microtubule and thus was being transported (Fig. S6). The kinesin-DNA run duration was 1.40 s, longer than the 1.04 s of motors alone (Fig. 2A). We interpret this longer duration to reflect the slower diffusion constant of the dsDNA relative to the motor alone, which enables motors to transiently detach and rebind before the DNA cargo has diffused away, thus extending the run duration 53. Notably, this slower diffusion constant should not play a role in the DNA tensiometer geometry because if the motor transiently detaches, it will be pulled backward by the elastic forces of the DNA and detected as a slip or detachment event.

Revised version: We carried out two additional control experiments. First, to confirm that the neutravidin used to link the DNA to the microtubule wasn’t affecting kinesin motility, we analyzed the run durations of kinesin-1 motors on neutravidin-coated microtubules and found no significant change in the unloaded run duration (Figure 2 – figure supplement 2). Second, we measured the run duration of kinesin-1 linked to a DNA tether that was not bound to the microtubule and thus was being transported (Figure 2 – figure supplement 2). The kinesin-DNA run duration was 1.40 s, longer than the 1.04 s of motors alone (Figure 2A). We interpret this longer duration to reflect the slower diffusion constant of the dsDNA relative to the motor alone, which enables motors to transiently detach and rebind before the DNA cargo has diffused away, thus extending the run duration 67. Notably, this slower diffusion constant should not play a role in the DNA tensiometer geometry because if the motor transiently detaches, it will be pulled backward by the elastic forces of the DNA and detected as a slip or detachment event.

Kinesin-3 detaches readily under low load

To determine whether sub-stall hindering loads affect motor detachment rates, we compared tensiometer ramp durations to the tensiometer stall and unloaded run durations (Figure 3). We defined ramp durations as the time the motor spends walking against the DNA spring before a slip or detachment, or before reaching stall. Although the dsDNA force-extension curve (Figure 1B) predicts negligible loads until the DNA is close to fully extended, there are still non-zero loads imposed during the ramp phase that may affect motor detachment. Based on 10-20% slower ramp velocities relative to unloaded velocities for each motor, we estimated the apparent force to be ∼1 pN (Table S2). To estimate the true detachment rate during the ramp phase in a way that takes into account both the observed detachments and ramps that successfully reach stall, we used a Markov process model, coupled with Bayesian inference methods (detailed in Supplementary Material) to estimate a duration parameter, τ, equivalent to the inverse of the detachment rate constant during a ramp. Each increment of time is considered to be an independent opportunity to detach while assuming a constant detachment rate; hence the probability of staying attached to the microtubule through a segment of duration Δ is e-Δ/τ. Using this method, ramp duration parameters, τ, were calculated for each motor, along with 95% credible regions. Finally, to allow for proper comparison, we performed a similar analysis to obtain the stall and unloaded duration parameters along with their 95% credible regions (Figure 3). As expected, the stall and unloaded durations were similar to estimates from curve fitting in Figure 2 (Table S1).

During Ramps, Kinesin-3 Detaches More Readily Than Under Zero Load.

Unloaded, ramp, and stall duration parameters were estimated using a Markov process model, coupled with Bayesian inference methods. Curves show the posterior probability distributions of the duration parameters for (A) kinesin-1, (B) kinesin-2 and (C) kinesin-3. Bars below each peak indicate the 95% credible regions for the ramp (green), unloaded (gray) and stall (blue) duration parameters. Notably, the estimated ramp durations are larger, the same, and smaller than the unloaded run durations for kinesin-1, -2, and -3, respectively. For the unloaded and stall durations, this estimation method produces almost identical values as the maximum likelihood estimates in Figure 2 (values provided in Table S1).

The prediction for a slip bond is that against the low loads experienced during ramps, the detachment rate should be equal to or faster than the unloaded detachment rate. This was the case for kinesin-2, where the ramp duration of 0.97 s was within 95% CI of the unloaded run duration of 1.08 s (Figure 3B, Table S1). In contrast, the kinesin-1 ramp duration of 2.49 s was much closer to the stall duration (3.05 s) than the unloaded run duration (1.05 s) (Figure 3A). One possible explanation for the longer kinesin-1 ramp is that the catch-bond character of kinesin-1 engages at low loads rather than rising proportionally to load or engaging only near stall.

Earlier version: The most notable ramp behavior was seen in kinesin-3, where the ramp duration of 0.75 s was nearly four-fold shorter than the unloaded run duration (2.76 s) and was more than two-fold shorter than the stall duration (1.90 s) (Figure 3C, Table S1). As expanded on in the Discussion, the positively charged ‘K-loop’ in the kinesin-3 motor KIF1A is known to interact electrostatically with the negatively charged C-terminal tail of tubulin 54-56; thus, it is reasonable that even the low loads imposed during ramps are sufficient to overcome these weak electrostatic interactions. The ramp duration is arguably the best definition of the time before KIF1A motors enter a partially dissociated ‘slip’ state, meaning that the observed unloaded durations represent a concatenation of multiple shorter runs interspersed by short diffusive events.

Revised version: The most notable ramp behavior was seen in kinesin-3, where the ramp duration of 0.75 s was nearly four-fold shorter than the unloaded run duration (2.76 s) and was more than two-fold shorter than the stall duration (1.90 s) (Figure 3C, Table S1). As expanded on in the Discussion, the positively charged ‘K-loop’ in the kinesin-3 motor KIF1A is known to interact electrostatically with the negatively charged C-terminal tail of tubulin 51,68,69; thus, it is reasonable that even the low loads imposed during ramps are sufficient to overcome these weak electrostatic interactions. The ramp duration is arguably the best definition of the time before KIF1A motors enter a partially dissociated ‘slip’ state, meaning that the observed unloaded durations represent a concatenation of multiple shorter runs interspersed by short diffusive events.

Earlier version: It was notable that the kinesin-3 stall durations at high load are longer than the ramp durations at low load, because this indicates that the kinesin-3 off-rate slows with increasing load. However, because kinesin-3 had the most slip events at stall, we were concerned that there may be undetected slip events below the 60 nm threshold of detection that led to an overestimation of the kinesin-3 stall duration. To test this hypothesis, we plotted the distribution of kinesin-3 slip distances at stall, fit an exponential, and calculated the fraction of missed slip events (Fig. S8). From this analysis, we calculated a correction factor of 1.42 that brought the kinesin-3 stall duration down 1.33 s. Notably, this stall duration value is still well above the kinesin-3 ramp duration value of 0.75 s in Fig. 3C and thus does not qualitatively change our conclusions.

Revised version: It was notable that the kinesin-3 stall durations at high load are longer than the ramp durations at low load, because this indicates that the kinesin-3 off-rate slows with increasing load. However, because kinesin-3 had the most slip events at stall, we were concerned that there may be undetected slip events below the 60 nm threshold of detection that led to an overestimation of the kinesin-3 stall duration. To test this hypothesis, we plotted the distribution of kinesin-3 slip distances at stall, fit an exponential, and calculated the fraction of missed slip events (Figure 4 – figure supplement 1). From this analysis, we calculated a correction factor of 1.42 that brought the kinesin-3 stall duration down 1.33 s. Notably, this stall duration value is still well above the kinesin-3 ramp duration value of 0.75 s in Figure 3C and thus does not qualitatively change our conclusions.

Motor reengagement kinetics vary between families

Earlier version: Stall plateaus were terminated by three types of events: 1) small slips that initiated a new ramp, typically within a single frame (∼40 ms), 2) the motor returning to the baseline and reengaging rapidly within a few frames (∼100 msec), or 3) the motor returning to the baseline for a few seconds before reengaging (Figure 4A). We defined a slip event as a displacement of >60 nm from the plateau (distinguishable from normal small fluctuations at stall; Figure 1E) that recovers before reaching within 400 nm of the baseline (outside the range of normal baseline fluctuations; Figure 1F and S2). These slip events have been observed previously for all three motor families in optical trapping experiments and are proposed to represent an intermediate state in which the motor exits the normal stepping cycle but remains associated with the microtubule 24,57-60. As an initial analysis, we quantified the fraction of events for each motor (Figure 4A) and found that kinesin-3 had the highest proportion of slip events while kinesin-2 had the lowest proportion. In the context of pulling a large cargo through the viscous cytoplasm or competing against dynein in a tug-of-war, these slip events enable the motor to continue generating force after a small rearward displacement, rather than fully detaching and ‘resetting’ to zero load. Thus, we reanalyzed the stall durations for the three motors where slips are not counted and only disengagements where the motor returns to the baseline are counted as stall termination events (Figure 4B & C, Figure S7). By this definition, stall durations were between 1.5 and 3-fold longer for each motor, including kinesin-3.

Revised version: Stall plateaus were terminated by three types of events: 1) small slips that initiated a new ramp, typically within a single frame (∼40 ms), 2) the motor returning to the baseline and reengaging rapidly within a few frames (∼100 msec), or 3) the motor returning to the baseline for a few seconds before reengaging (Figure 4A). We defined a slip event as a displacement of >60 nm from the plateau (distinguishable from normal small fluctuations at stall; Figure 1E) that recovers before reaching within 400 nm of the baseline (outside the range of normal baseline fluctuations; Figure 1F and S2). These slip events have been observed previously for all three motor families in optical trapping experiments and are proposed to represent an intermediate state in which the motor exits the normal stepping cycle but remains associated with the microtubule 24,70–73. As an initial analysis, we quantified the fraction of events for each motor (Figure 4A) and found that kinesin-3 had the highest proportion of slip events while kinesin-2 had the lowest proportion. In the context of pulling a large cargo through the viscous cytoplasm or competing against dynein in a tug-of-war, these slip events enable the motor to continue generating force after a small rearward displacement, rather than fully detaching and ‘resetting’ to zero load. Thus, we reanalyzed the stall durations for the three motors where slips are not counted and only disengagements where the motor returns to the baseline are counted as stall termination events (Figure 4B & C, Figure 4 – figure supplement 2). By this definition, stall durations were between 1.5 and 3-fold longer for each motor, including kinesin-3.

Restart Kinetics for Kinesin-1, -2 and -3.

Earlier version: (A) Fraction of slip, fast rebinding, and slow rebinding events for each motor, with example kymographs for each (top; scale bars are 0.5 μm and 0.2 s). Solid colors indicate slips during stall, where the motor resumes a new ramp within a single frame (∼40 ms), crosshatching indicates rapid reattachment events (100 ms) following fall to baseline, and open bars indicate slow reattachment events with >100 ms fluctuations around baseline. (B) Kinesin-3 stall durations, with unloaded run times in gray, stall durations terminated by slips in light green, and stall durations terminated by falling to the baseline (ignoring slips) in dark green. Unloaded and stall durations (replotted from Figure 2) were fit with single exponential functions in MEMLET. Stall durations ignoring slips were fit with a bi-exponential by least squares (ρ1 = 2.01 s [95% CI: 1.52, 2.35 s], A1 = 0.66 s [0.54, 0.83 s], ρ2 = 11.0 s [9.18, 13.70], A2 = 0.33 [0.23, 0.48]). Weighted average of the two time constants is displayed in plot for comparison to other time constants. Similar results for kinesin-1 and -2 are shown in Figure S7. (C) Comparison of stall durations for kinesins -1, -2 and -3 with slips observed as stall terminations or ignored. (D-F) Distribution of restart times for each motor fit to a tri-exponential (least squares). Confidence intervals of parameters determined by bootstrapping with 1000 iterations are given in Table S3.

Revised version: (A) Fraction of slip, fast rebinding, and slow rebinding events for each motor, with example kymographs for each (top; scale bars are 0.5 μm and 0.2 s). Solid colors indicate slips during stall, where the motor resumes a new ramp within a single frame (∼40 ms), crosshatching indicates rapid reattachment events (100 ms) following fall to baseline, and open bars indicate slow reattachment events with >100 ms fluctuations around baseline. (B) Kinesin-3 stall durations, with unloaded run times in gray, stall durations terminated by slips in light green, and stall durations terminated by falling to the baseline (ignoring slips) in dark green. Unloaded and stall durations (replotted from Figure 2) were fit with single exponential functions in MEMLET. Stall durations ignoring slips were fit with a bi-exponential by least squares (ρ1 = 2.01 s [95% CI: 1.52, 2.35 s], A1 = 0.66 s [0.54, 0.83 s], ρ2 = 11.0 s [9.18, 13.70], A2 = 0.33 [0.23, 0.48]). Weighted average of the two time constants is displayed in plot for comparison to other time constants. Similar results for kinesin-1 and -2 are shown in Figure 4 – figure supplement 2. (C) Comparison of stall durations for kinesins -1, -2 and -3 with slips observed as stall terminations or ignored. (D-F) Distribution of restart times for each motor fit to a tri-exponential (least squares). Confidence intervals of parameters determined by bootstrapping with 1000 iterations are given in Table S3.

Earlier version: To obtain a more complete picture of the motor reengagement kinetics for each motor, we plotted the time before starting a new ramp (trestarting), including all slips and reattachments (Figure 4D-F). In each case, the distributions included a fast phase and a long tail of slower events. The distributions were fit with a tri-exponential function with the fast phase (30, 130 and 40 msec, respectively) accounting for roughly half of the events (Figure 4 D-F, Table S3). The fast population corresponds to slip and fast reattachment events classified in Figure 4A. The slower phases, which represent detachment events where the motor fluctuated around the baseline before initiating a new ramp, were the fastest for kinesin-3 (time constants of 0.15 and 1.34 s) and the slowest for kinesin-1 (time constants of 1.51 and 20.9 s). Interestingly, the order of the reattachment kinetics (kinesin-3 > kinesin-2 > kinesin-1) and the ∼10-fold ratio of kinesin-3 to kinesin-1 match published bimolecular on-rate constants for microtubule binding from stopped-flow experiments (1.1, 4.6, and 17 μM-1 s-1 for kinesin-1,-2 and - 3, respectively 56,61,62).

Revised version: To obtain a more complete picture of the motor reengagement kinetics for each motor, we plotted the time before starting a new ramp (trestarting), including all slips and reattachments (Figure 4D-F). In each case, the distributions included a fast phase and a long tail of slower events. The distributions were fit with a tri-exponential function with the fast phase (30, 130 and 40 msec, respectively) accounting for roughly half of the events (Figure 4 D-F, Table S3). The fast population corresponds to slip and fast reattachment events classified in Figure 4A. The slower phases, which represent detachment events where the motor fluctuated around the baseline before initiating a new ramp, were the fastest for kinesin-3 (time constants of 0.15 and 1.34 s) and the slowest for kinesin-1 (time constants of 1.51 and 20.9 s). Interestingly, the order of the reattachment kinetics (kinesin-3 > kinesin-2 > kinesin-1) and the ∼10-fold ratio of kinesin-3 to kinesin-1 match published bimolecular on-rate constants for microtubule binding from stopped-flow experiments (1.1, 4.6, and 17 μM-1 s-1 for kinesin-1,-2 and - 3, respectively 50,51,74).

Simulating Potential Catch-Bond Mechanisms

Earlier version: To compare the motor detachment and reattachment kinetics between the three kinesin transport families, we carried out stochastic simulations of load-dependent motor stepping, unbinding and rebinding. For simplicity, we reduced the chemomechanical cycle down to a single strongly-bound state (ATP and nucleotide-free states) and a single weakly-bound state (ADP and ADP-Pi states) (Figure 5A). Based on the published load and ATP dependencies of substeps in the kinesin-1 chemomechanical cycle 57, we incorporated a load-dependent strong-to-weak transition, ks-w. Based on our restarting durations from Figure 4 and previous work 57-60, we included both a slip state, from which the motor recovers rapidly, and a detached state associated with a slower recovery. Runs or stalls are terminated by transition from the weakly-bound state into the slip state (kslip). Based on backstepping rates observed in optical tweezer experiments, we incorporated a load-independent backward stepping rate of 3 s-1, meaning that stall is defined as the load at which forward stepping slows to 3 s-1 63,64. For simplicity, we set ks-w = kw-s at zero load, meaning that the motor spends half of its cycle in each state, 65 with the rates set to match the unloaded velocity for each motor.

Revised version: To compare the motor detachment and reattachment kinetics between the three kinesin transport families, we carried out stochastic simulations of load-dependent motor stepping, unbinding and rebinding. For simplicity, we reduced the chemomechanical cycle down to a single strongly-bound state (ATP and nucleotide-free states) and a single weakly-bound state (ADP and ADP-Pi states) (Figure 5A). Based on the published load and ATP dependencies of substeps in the kinesin-1 chemomechanical cycle 70, we incorporated a load-dependent strong-to-weak transition, ks-w. Based on our restarting durations from Figure 4 and previous work 70–73, we included both a slip state, from which the motor recovers rapidly, and a detached state associated with a slower recovery. Runs or stalls are terminated by transition from the weakly-bound state into the slip state (kslip). Based on backstepping rates observed in optical tweezer experiments, we incorporated a load-independent backward stepping rate of 3 s-1, meaning that stall is defined as the load at which forward stepping slows to 3 s-1 75,76. For simplicity, we set ks-w = kw-s at zero load, meaning that the motor spends half of its cycle in each state, 77 with the rates set to match the unloaded velocity for each motor.

Chemomechanical Model of Proposed Catch-bond Mechanism.

Earlier version: (A) Diagram of kinesin chemomechanical cycle model consisting of strongly- and weakly-bound states that make up the stepping cycle, and slip and detached states that terminate runs and stalls. Note that two pathways of detachment from the slip state (and reattachment) are incorporated into the model, but only one pathway is shown for simplicity (see Supplementary Methods for details). (B) Table of rate constants used to simulate unloaded and stall durations and restarting times. All rate constants are derived from fits to experimental data, as described in Supplemental Methods. kS-W and kslip depended exponentially on load with δ for kS-W of -2.7, -2.4, and -3.6 nm and δ for kslip of 1.6, 1.3 and 2.7 nm for kinesin-1, -2 and -3, respectively; see also Figure S9A). (C-E) Experimental (symbols) and simulated (lines) unloaded and stall durations. 10,000 events were simulated for each condition and plotted with minimum cutoffs matching experiments. Kinesin-3 ramp durations were taken from parameter estimated in Figure 3. (F-H) Experimental (symbols) and simulated (lines) restart times.

Revised version: (A) Diagram of kinesin chemomechanical cycle model consisting of strongly- and weakly-bound states that make up the stepping cycle, and slip and detached states that terminate runs and stalls. Note that two pathways of detachment from the slip state (and reattachment) are incorporated into the model, but only one pathway is shown for simplicity (see Supplementary Methods for details). (B) Table of rate constants used to simulate unloaded and stall durations and restarting times. All rate constants are derived from fits to experimental data, as described in Supplemental Methods. kS-W and kslip depended exponentially on load (k(F)=k0eFδkBT) with δ for kS-W of -2.7, -2.4, and -3.6 nm and δ for kslip of 1.6, 1.3 and 2.7 nm for kinesin-1, -2 and -3, respectively; see also Figure 5 – figure supplement 1A). (C-E) Experimental (symbols) and simulated (lines) unloaded and stall durations. 10,000 events were simulated for each condition and plotted with minimum cutoffs matching experiments. Kinesin-3 ramp durations were taken from parameter estimated in Figure 3. (F-H) Experimental (symbols) and simulated (lines) restart times.

Using parameters chosen to match motor behavior under no load (Figure 5B), we were able to reproduce the unloaded run durations for all three motors, as expected (Figure 5C-E). Next, to match the experimental stall durations we incorporated a negative load dependence into the strong-to-weak transition rate (so that stepping slows with load) and a positive load dependence into the transition from the weakly-bound state into the slip state (so slipping occurs more frequently at higher loads). With these parameters (Figure 5B), we were able to reproduce the stall duration distribution for all three motors (Figure 5C-E). Importantly, in this model formulation, dissociation from the weakly-bound state acts as a slip-bond and the kinesin catch-bond characteristics are achieved by the motor spending a larger fraction of its cycle in the strongly-bound state under increasing loads.

Earlier version: To complete our model, we simulated recovery from the slip and detached states. The rate of rescue from the slip state (kresc) was set based on the duration and relative amplitude of the fast phase of the restarting times in Figure 4. We posited that detachment follows the slip state, consistent with previous formulations 57,60. The two slower restarting time constants from Figure 4 were used to set the reattachment rates, kreatt (see Supplementary Methods for details). Using this approach, we were able to reproduce the restarting durations for all three motors (Figure 5F-H).

Revised version: To complete our model, we simulated recovery from the slip and detached states. The rate of rescue from the slip state (kresc) was set based on the duration and relative amplitude of the fast phase of the restarting times in Figure 4. We posited that detachment follows the slip state, consistent with previous formulations 70,73. The two slower restarting time constants from Figure 4 were used to set the reattachment rates, kreatt (see Supplementary Methods for details). Using this approach, we were able to reproduce the restarting durations for all three motors (Figure 5F-H). We note that recapitulating the triexponential restart time distribution in Figure 4D-F required this slip/detached formulation and that lumping all events into a single detached state resulted in single-exponential distribution of recovery times.

Earlier version: The rate constants derived from this model allow for a comparison of the specific transitions that differ between the three motor families. Three features are notable. First, for kinesin-3 the transition into the slip state at stall, kslip, is the fastest of the three families, consistent with the observation by eye of the plateaus (Figure 1H and 4A) and consistent with previous three-bead optical trapping results 24. Second, the rates of rescue from the slip state, kresc, for the three motors match within a factor of two. Thus, the fast bimolecular on-rates from stopped flow and the relatively short durations before restarting we observe for kinesin-3 (Figure 4) do not result from a faster reengagement out of the slip state. Third, the slow reattachment rate for kinesin-3 is 10-fold faster than for kinesin-1 and -2. Hence, in this model formulation, both the fast bimolecular on-rates in solution56 and short restart durations observed here for kinesin-3 result from recovery from a detached state rather than rescue from a slip state.

Revised version: The rate constants derived from this model allow for a comparison of the specific transitions that differ between the three motor families. Three features are notable. First, for kinesin-3 the transition into the slip state at stall, kslip, is the fastest of the three families, consistent with the observation by eye of the shorter plateaus for kinesin-3 (Figure 1H and 4A) and consistent with slips observed in a previous three-bead optical trapping study 24. Second, the rates of rescue from the slip state, kresc, for the three motors match within a factor of two. Thus, the fast bimolecular on-rates from stopped flow and the relatively short durations before restarting we observe for kinesin-3 (Figure 4) do not result from a faster reengagement out of the slip state in this model. Third, the slow reattachment rate for kinesin-3 is 10-fold faster than for kinesin-1 and -2. Hence, in this model formulation, both the fast bimolecular on-rates in solution 51 and short restart durations observed here for kinesin-3 result from recovery from a detached state rather than rescue from a slip state.

Discussion

Earlier version: Understanding how motors respond vectorially to external loads is crucial for understanding cargo transport in complex intracellular geometries and how kinesin motors compete against dynein in bidirectional transport 13,19. Optical tweezer experiments have provided many essential details of the kinesin mechanochemical cycle under load; however, the bead diameters needed to achieve substantial trapping forces impose vertical forces on the motors. Using DNA as a nanospring enables mechanical experiments using a standard TIRF microscope and allows for simultaneous monitoring of numerous motor-DNA complexes in a single imaging field. With this geometry, a kinesin motor pulls against the elastic force of a stretched DNA nearly parallel to the microtubule, matching the geometry of vesicles measuring a few tens of nanometers. Similar approaches have been used to study myosin, dynein and kinesin-1 in both single-molecule and gliding assays 33-35,37,38,66,67. The most striking observation was that members of all three kinesin transport families show catch-bond behavior in which off-rates at stall are slower than those at low or zero loads. Additionally, following disengagement from the microtubule, the three motor families reengaged with the microtubule with complex and family-specific kinetics.

Revised version: Understanding how motors respond vectorially to external loads is crucial for understanding cargo transport in complex intracellular geometries and how kinesin motors compete against dynein in bidirectional transport 13,19. Optical tweezer experiments have provided many essential details of the kinesin mechanochemical cycle under load; however, the bead diameters needed to achieve substantial trapping forces impose vertical forces on the motors. Using DNA as a nanospring enables mechanical experiments using a standard TIRF microscope and allows for simultaneous monitoring of numerous motor-DNA complexes in a single imaging field. With this geometry, a kinesin motor pulls against the elastic force of a stretched DNA nearly parallel to the microtubule, matching the geometry of vesicles measuring a few tens of nanometers. Similar approaches have been used to study myosin, dynein and kinesin-1 in both single-molecule and gliding assays 33–35,37,38,78,79. The most striking observation was that members of all three kinesin transport families show catch-bond behavior in which off-rates at stall are slower than those at low or zero loads. Additionally, following disengagement from the microtubule, the three motor families reengaged with the microtubule with complex and family-specific kinetics.

Comparison to previous work

Earlier version: Despite the clear slip-bond behavior of kinesin-1 seen in single-bead optical traps, there has been growing evidence that kinesin-1 detachment is sensitive to the direction of load. Motor engagement times were shown to decrease when larger beads were employed in single-bead traps, and to be extended in geometries that minimize vertical forces such as the three-bead geometry or when a DNA tether was used 32,68. In a study that employed DNA-tethered kinesin-1 to extract tubulin the microtubule lattice, pulling durations of ∼30 s were observed at the lowest motor concentrations, indicative of kinesin-1 catch-bond behavior 33. When kinesin-1 was connected to micron-scale beads through a DNA linker and hydrodynamic forces parallel to the microtubule imposed, dissociation rates were relatively insensitive to loads up to ∼3 pN, inconsistent with slip-bond characteristics and instead characteristic of an ideal-bond 34. The 3 s kinesin-1 stall duration in our tensiometer falls between the 30 s value for tubulin pulling and the 1.3 – 1.5 s engagement times measured in optical trap and hydrodynamic assays where vertical forces are minimized 32-34,68. In contrast to kinesin-1, kinesin-3 (KIF1A) median engagement durations were found to be similarly short in both the one-bead (69 ms) and three bead (62 ms) optical trap geometries 24, much shorter than our 1.9 s stall duration in the DNA tensiometer. One difference may be that the millisecond temporal resolution in the optical tweezer enabled detection of many small slips that were undetected by our fluorescence approach; however, extrapolation of our measured kinesin-3 slip distances below our 60 nm detection limit (Fig. S8) argues against this. Secondly, the microtubule is held under tension in the three-bead experiment, which may alter the lattice properties. Finally, it can’t be ruled out that KIF1A is particularly sensitive to vertical loads and small lateral or vertical forces present in the three-bead geometry are absent in the DNA tensiometer geometry. To date there have been no studies on kinesin-2 where vertical loads are minimized, and in single-bead optical traps the off-rate depends strongly on load69.

Revised version: Despite the clear slip-bond behavior of kinesin-1 seen in single-bead optical traps, there has been growing evidence that kinesin-1 detachment is sensitive to the direction of load. Motor engagement times were shown to decrease when larger beads were employed in single-bead traps, and to be extended in geometries that minimize vertical forces such as the three-bead geometry or when a DNA tether was used 32,80. In a study that employed DNA-tethered kinesin-1 to extract tubulin the microtubule lattice, pulling durations of ∼30 s were observed at the lowest motor concentrations, indicative of kinesin-1 catch-bond behavior 33. When kinesin-1 was connected to micron-scale beads through a DNA linker and hydrodynamic forces parallel to the microtubule imposed, dissociation rates were relatively insensitive to loads up to ∼3 pN, inconsistent with slip-bond characteristics and instead characteristic of an ideal-bond 34. The 3 s kinesin-1 stall duration in our tensiometer falls between the 30 s value for tubulin pulling and the 1.3 – 1.5 s engagement times measured in optical trap and hydrodynamic assays where vertical forces are minimized 32–34,80. In contrast to kinesin-1, kinesin-3 (KIF1A) median engagement durations were found to be similarly short in both the one-bead (69 ms) and three bead (62 ms) optical trap geometries 24, much shorter than our 1.9 s stall duration in the DNA tensiometer. One difference may be that the millisecond temporal resolution in the optical tweezer enabled detection of many small slips that were undetected by our fluorescence approach; however, extrapolation of our measured kinesin-3 slip distances below our 60 nm detection limit (Figure 4 – figure supplement 1) argues against this. Stiffness differences are an unlikely explanation because at stall the stiffness of the DNA tether (∼4 fold stiffer than optical tweezer) is still sufficiently low to allow for dynamic motor stepping at stall, and in any case it is still below the estimated motor stiffness (see Geometry Calculations in Supplementary methods). One potential explanation is that the microtubule is held under tension in the three-bead experiment, which may alter the lattice properties and affect motor interactions. Finally, it can’t be ruled out that KIF1A is particularly sensitive to vertical loads and small lateral or vertical forces present in the three-bead geometry are absent in the DNA tensiometer geometry. To date there have been no studies on kinesin-2 where vertical loads are minimized, and in single-bead optical traps the off-rate depends strongly on load 52.

Transport kinesins have a catch-bond behavior under hindering loads

Earlier version: What is the mechanism of the observed catch-bond behavior? Cell adhesion proteins such as integrins, selectins, and FimH have been shown to form longer lasting bonds under load, with the proposed mechanisms generally involving an allosteric effect that strengthens the protein:protein interface 70-72. However, motor proteins are different in that they cycle in a nucleotide-dependent way between strongly- and weakly bound states, offering multiple potential mechanisms for slower dissociation rates under load. For instance, under a few piconewtons of load, Myosin I was shown to dissociate nearly two orders of magnitude slower than in the absence of load, an effect attributed to load-dependent trapping of ADP in the active site that maintained the motor a high-affinity binding state 73. Dynein was also shown to have catch-bond behavior over certain ranges of resisting loads, though the precise mechanism is unclear 43,44,74,75.

Revised version: What is the mechanism of the observed catch-bond behavior? Cell adhesion proteins such as integrins, selectins, and FimH have been shown to form longer lasting bonds under load, with the proposed mechanisms generally involving an allosteric effect that strengthens the protein: protein interface 81–83. However, motor proteins are different in that they cycle in a nucleotide-dependent way between strongly- and weakly bound states, offering multiple potential mechanisms for slower dissociation rates under load. For instance, under a few piconewtons of load, Myosin I was shown to dissociate nearly two orders of magnitude slower than in the absence of load, an effect attributed to load-dependent trapping of ADP in the active site that maintained the motor a high-affinity binding state 84. Dynein was also shown to have catch-bond behavior over certain ranges of resisting loads, though the precise mechanism is unclear 56,57,60,85,86.

We interpreted our stepping, detachment, and reattachment results using a model that incorporates a load-dependent strong-to-weak transition and a load-dependent entry into a transient ‘slip’ state preceding detachment. The key feature of the model is that under load the motor spends an increasing fraction of its hydrolysis cycle in a strongly-bound state that resists dissociation.

The role of vertical forces in motor detachment

Earlier version: We next asked whether by considering the different geometries we could reconcile our catch-bond observations with previous single-bead optical tweezer kinesin-1 slip-bond measurements that found the kinesin-1 off-rate increased from 1.11 s-1 at zero load to 2.67 s-1 at 6 pN 25. Using a 440 nm bead diameter and estimated motor length of 35 nm that results in the force being imposed on the motor at a 60° angle (Figure S9B), a 6 pN stall force parallel to the microtubule corresponds to a 10 pN force perpendicular to the microtubule 25,31. A model developed by Khataee and Howard was able to fully account for these geometry-dependent off-rates using a two-step detachment process having catch-bond behavior for parallel loads and slip-bond behavior for vertical loads 31. Applying that model to our geometry and assuming a purely horizontal load and a 6 pN stall force, the predicted stall duration for kinesin-1 is 77 s, much longer than the 3 s we measure (Figure S9). We approached the detachment process differently, by incorporating a load-dependent transition in the hydrolysis cycle (kS-W) and a load-dependent exit from the hydrolysis cycle (kslip). To explore effects of vertical forces using our two-state model, we incorporated both horizontal and vertical loads as accelerating detachment from the weakly-bound state, as follows: Here F|| and F⊥ are the magnitude of the parallel and perpendicular loads and δ|| and δ⊥ represent the distance parameters in each direction (Figure S9). Using δ|| = 1.61 nm (Figure 5B), we found that by setting δ⊥ = 1.58 nm, we were able to reproduce the slip bond behavior observed in the single-bead optical trap experiments (Figure S9). Notably, this model implies that vertical and horizontal forces have similar effects on the transition rate into the slip state. We stress that this model is a hypothesis that needs further testing. Nonetheless, this is a simple formulation that shows that a motor can display either catch bond or slip bond behavior depending on the geometry of the imposed loads.

Revised version: We next asked whether by considering the different geometries we could reconcile our catch-bond observations with previous single-bead optical tweezer kinesin-1 slip-bond measurements that found the kinesin-1 off-rate increased from 1.11 s-1 at zero load to 2.67 s-1 at 6 pN 25. Using a 440 nm bead diameter and estimated motor length of 35 nm that results in the force being imposed on the motor at a 60° angle (Figure 5 – figure supplement 1B), a 6 pN stall force parallel to the microtubule corresponds to a 10 pN force perpendicular to the microtubule 25,31. A model developed by Khataee and Howard was able to fully account for these geometry-dependent off-rates using a two-step detachment process having catch-bond behavior for parallel loads and slip-bond behavior for vertical loads 31. Applying that model to our geometry and assuming a purely horizontal load and a 6 pN stall force, the predicted stall duration for kinesin-1 is 77 s, much longer than the 3 s we measure (Figure 5 – figure supplement 1). We approached the detachment process differently, by incorporating a load-dependent transition in the hydrolysis cycle (kS-W) and a load-dependent exit from the hydrolysis cycle (kslip). To explore effects of vertical forces using our two-state model, we incorporated both horizontal and vertical loads as accelerating detachment from the weakly-bound state, as follows: kslip(F)=kslip0e(F‖δ‖+F⊥δ⊥)kBT Here F|| and F⊥ are the magnitude of the parallel and perpendicular loads and δ|| and δ⊥ represent the distance parameters in each direction (Figure 5 – figure supplement 1). Using δ|| = 1.61 nm (Figure 5B), we found that by setting δ⊥ = 1.58 nm, we were able to reproduce the slip bond behavior observed in the single-bead optical trap experiments (Figure 5 – figure supplement 1). Notably, this model implies that vertical and horizontal forces have similar effects on the transition rate into the slip state. We stress that this model is a hypothesis that needs further testing. Nonetheless, this is a simple formulation that shows that a motor can display either catch bond or slip bond behavior depending on the geometry of the imposed loads.

Ramps reveal detachment behaviors at low loads

Earlier version: In addition to reporting on the detachment properties at stall, our DNA tensiometer provides new insights into fast rebinding that occurs during unloaded runs of kinesin-3. It has long been appreciated that the kinesin-3 motor KIF1A achieves long run lengths due to electrostatic attraction between its positively charged Loop-12 (K-loop) and the negatively charged C-terminal tail of tubulin 55,56,76-81. Furthermore, the KIF1A off-rate in ADP, in which the motor diffuses on the microtubule lattice, was found to match the off-rate during processive stepping in ATP 56. The relatively high microtubule affinity of this weakly-bound state suggests that the motor may be undergoing diffusive episodes between processive runs, while maintaining association with the microtubule.

Revised version: In addition to reporting on the detachment properties at stall, our DNA tensiometer provides new insights into fast rebinding that occurs during unloaded runs of kinesin-3. It has long been appreciated that the kinesin-3 motor KIF1A achieves long run lengths due to electrostatic attraction between its positively charged Loop-12 (K-loop) and the negatively charged C-terminal tail of tubulin 45,51,69,87–91. Furthermore, the KIF1A off-rate in ADP, in which the motor diffuses on the microtubule lattice, was found to match the off-rate during processive stepping in ATP 51. The relatively high microtubule affinity of this weakly-bound state suggests that the motor may be undergoing diffusive episodes between processive runs, while maintaining association with the microtubule.

Earlier version: Our DNA tensiometer offers a way to test the hypothesis that the long, unloaded run lengths of KIF1A are due to a concatenation of shorter runs. Due to the nonlinearity of the dsDNA force-extension curve in our DNA tensiometer, the motor is walking against forces below 1 pN for roughly 90% of the distance to stall (Figure 1B). Consistent with this, motor velocities before stall were nearly constant (Figure 1D and S3), and averaged ∼15% slower than unloaded velocities (Table S2), which corresponds to ∼1 pN of force if the force-velocity relationship is linear. Using a Bayesian Inference approach that takes into account motors that dissociate during the ramps as well as those that complete ramps by achieving stall, we measured a nearly four-fold faster KIF1A detachment rate during ramps than under zero load (Figure 3). If, under zero load, the long runs observed were actually a concatenation of a series of shorter runs connected by diffusive weakly-bound events, the diffusive state would likely be unable to withstand even the sub-pN forces from the DNA spring 80. For instance, based on a 0.044 μm2/s diffusion coefficient (equivalent to a ∼0.1 pN-s/μm drag coefficient 80,82), if the motor were in a weakly-bound state for 10 ms, a 1 pN force would pull the motor back 100 nm. Thus, in considering whether KIF1A acts as a catch bond, we used this ramp duration as the best approximation for the true run unloaded length in the absence of diffusive events.

Revised version: Our DNA tensiometer offers a way to test the hypothesis that the long, unloaded run lengths of KIF1A are due to a concatenation of shorter runs. Due to the nonlinearity of the dsDNA force-extension curve in our DNA tensiometer, the motor is walking against forces below 1 pN for roughly 90% of the distance to stall (Figure 1B). Consistent with this, motor velocities before stall were nearly constant (Figure 1D and S3), and averaged ∼15% slower than unloaded velocities (Table S2), which corresponds to ∼1 pN of force if the force-velocity relationship is linear. Using a Bayesian Inference approach that takes into account motors that dissociate during the ramps as well as those that complete ramps by achieving stall, we measured a nearly four-fold faster KIF1A detachment rate during ramps than under zero load (Figure 3). If, under zero load, the long runs observed were actually a concatenation of a series of shorter runs connected by diffusive weakly-bound events, the diffusive state would likely be unable to withstand even the sub-pN forces from the DNA spring 90. For instance, based on a 0.044 μm2/s diffusion coefficient (equivalent to a ∼0.1 pN-s/μm drag coefficient 90,92), if the motor were in a weakly-bound state for 10 ms, a 1 pN force would pull the motor back 100 nm. Thus, in considering whether KIF1A acts as a catch bond, we used this ramp duration of 0.75 s (Table S1) as the best approximation for the true run unloaded length in the absence of diffusive events.

The ramp durations of kinesin-1 and kinesin-2 also provide insights into how load alters their interactions with microtubules. For kinesin-2, the predicted ramp duration was not statistically different from the unloaded run duration, suggesting that unloaded runs do not include short diffusive episodes. Interestingly the predicted ramp duration for kinesin-1 was nearly the same as the stall duration and much longer than the unloaded duration. One possibility is that the catch-bond effect of hindering load comes into play at low loads and not only at stall where the motor has slowed considerably.

Motor slips and detachments reflect different processes

Earlier version: Because the DNA tensiometer tethers the motor near the microtubule, such that repeated binding and unbinding events occur, it enables comparison of family-dependent differences in kinesin rebinding kinetics. In addition to clear detachment events, rapid slip and recovery events were observed for all three motors, with highest frequency for kinesin-3 and lowest frequency for kinesin-1. Backwards slipping while maintaining association with the microtubule was first seen for kinesin-8 motors, which are highly processive yet generate only small forces 58. Similar backward slips at stall were observed for kinesin-1, with kinetics that suggested a transition such as phosphate release precedes dissociation60. Subsequent higher resolution work, enabled by small Germanium nanoparticles, revealed a staircase pattern during these slips with ∼8 nm steps of mean duration 73 μsec, suggesting that the motor was transiently interacting with each tubulin subunit as it slipped backward. Similar slips have also been observed for kinesin-2 and two kinesin-3 family members, KIF1A and KIF1C 24,59.

Revised version: Because the DNA tensiometer tethers the motor near the microtubule, such that repeated binding and unbinding events occur, it enables comparison of family-dependent differences in kinesin rebinding kinetics. In addition to clear detachment events, rapid slip and recovery events were observed for all three motors, with highest frequency for kinesin-3 and lowest frequency for kinesin-1. Backwards slipping while maintaining association with the microtubule was first seen for kinesin-8 motors, which are highly processive yet generate only small forces 71. Similar backward slips at stall were observed for kinesin-1, with kinetics that suggested a transition such as phosphate release precedes dissociation73. Subsequent higher resolution work, enabled by small Germanium nanoparticles, revealed a staircase pattern during these slips with ∼8 nm steps of mean duration 73 μsec, suggesting that the motor was transiently interacting with each tubulin subunit as it slipped backward. Similar slips have also been observed for kinesin-2 and two kinesin-3 family members, KIF1A and KIF1C 24,72.

Earlier version: There is some dispute in the literature regarding the kinetics of kinesin-1 recovery from the slip state. Using a single-bead trap, Sudhakar found that 80% of restart events were slips with a time constant of 128 ms 57, whereas Toleikis measured slip recoveries that were essentially at the limit of detection (∼1 ms) 60. Using a three-bead trap, Pyrpassopoulis measured a 10 ms slip time constant for kinesin-1 32. Our 40 ms slip time constant, which accounts for half of recovery events, is limited by the camera frame rate, and thus is likely an overestimate. In the three-bead geometry, kinesin-3 (KIF1A) slips recovered with a time constant of 1 ms 24, faster than the 30 msec (upper limit estimate) we observe. Thus, the precise recovery rate is dependent on the detailed measurement and analysis used. In our DNA tensiometer results, the higher frequency of slips for kinesin-3 relative to kinesin-1 is seen by direct counting (Figure 4A), by the large enhanced stall duration when slips are not counted as termination events (Figure 4C), and by the fast kslip parameter under load in the kinesin-3 model (Figure 5B).

Revised version: There is some dispute in the literature regarding the kinetics of kinesin-1 recovery from the slip state. Using a single-bead trap, Sudhakar found that 80% of restart events were slips with a time constant of 128 ms 70, whereas Toleikis measured slip recoveries that were essentially at the limit of detection (∼1 ms) 73. Using a three-bead trap, Pyrpassopoulis measured a 10 ms slip time constant for kinesin-1 32. Our 40 ms slip time constant, which accounts for half of recovery events, is limited by the camera frame rate, and thus is likely an overestimate. In the three-bead geometry, kinesin-3 (KIF1A) slips recovered with a time constant of 1 ms 24, faster than the 30 msec (upper limit estimate) we observe. Thus, the precise recovery rate is dependent on the detailed measurement and analysis used. In our DNA tensiometer results, the higher frequency of slips for kinesin-3 relative to kinesin-1 is seen by direct counting (Figure 4A), by the large enhanced stall duration when slips are not counted as termination events (Figure 4C), and by the fast kslip parameter under load in the kinesin-3 model (Figure 5B).

Catch-bond behavior provides insights into tug-of-war with dynein

Earlier version: In previous simulations of kinesin-dynein bidirectional transport, we found that the strongest determinants of kinesin’s ability to compete against dynein were the load-dependent motor dissociation rate and the motor rebinding rate 22,23. Simply put, if motors detach, then the opposing motor wins. The finding here that all three dominant kinesin transport families display catch-bond like behavior at stall necessitates a reevaluation of how motors function during a tug-of-war. There is evidence that dynein forms a catch bond or at least an ideal (load independent) bond 14,43,44,74,75,83; thus, kinesins and dyneins are primed to strongly oppose one another.

Revised version: In previous simulations of kinesin-dynein bidirectional transport, we found that the strongest determinants of kinesin’s ability to compete against dynein were the load-dependent motor dissociation rate and the motor rebinding rate 22,23. Simply put, if motors detach, then the opposing motor wins. The finding here that all three dominant kinesin transport families display catch-bond like behavior at stall necessitates a reevaluation of how motors function during a tug-of-war. There is evidence that dynein forms a catch bond or at least an ideal (load independent) bond 14,56,57,60,85,86,93; thus, kinesins and dyneins are primed to strongly oppose one another.

The catch bond results here help to explain previous in vitro work in which one kinesin and one dynein were connected through a complementary ssDNA 14,28,30. It was found that the motor pairs had periods of near-zero velocity that lasted for many seconds, considerably longer than kinesin’s unloaded off-rate. Furthermore, kinesin-2 and kinesin-3 also showed these sustained slow tug-of-war periods despite their reported faster load-dependent off-rates from optical tweezer studies. The functional catch-bond behavior observed here provides a simple explanation for these sustained kinesin-dynein stalemates.

Importantly, the load-dependent off-rates of both kinesins and dynein are expected to depend on the cargo geometry. A 30 nm vesicle would lead to forces on the motor nearly parallel to the microtubule surface, whereas when transporting a micron-scale mitochondria the vertical forces would be larger than the horizontal forces. Cargo geometry and stiffness are also expected to play a role; for instance, deformation of a cargo, either due to compliance of the cargo or to multiple motors pulling on it will tend to reduce vertical force components on the motors. Additionally, any ‘rolling’ of a spherical cargo following motor detachment will tend to suppress the motor reattachment rate. The present work emphasizes that along with motor type, motor number, motor autoinhibition, and the growing list of regulatory proteins, the geometry with kinesin and dynein engage in a tug-of-war can be an important determinant of the speed and direction of cargo transport in cells.

Methods

DNA Tensiometer Construction

Earlier version: For the dsDNA spring, a 5’ biotinylated forward primer (5’-/5Biosg/TGC CTC CGT GTA AGG GGG AT-3’) and a reverse primer with a 5’ overhang (5’-/GGG CCA TCG CCA ATT GGA GTA /idSp/ GTG AGT TAA AGT TGT ACT CGA GTT TGT GTC CAA GAA -3’) were used to create a 3009 bp dsDNA by PCR from plasmid Mus Musculus BicD2-sf-GFP (aa 25-425) in pet28a. The abasic Int 1’,2’-Dideoxyribose spacer (idSp) creates an overhang by terminating the polymerase. All oligonucleotides were purchased from IDT. Each 50 μL PCR reaction contained: 1x Phusion HF buffer, 198 μM dNTPs, 2 μM dCTP-Cy5, 0.5 μM primers, 3 ng template DNA and 1 U/50 μL HF Phusion Polymerase. Fluorescent dsDNA used in Figure S4 had 10 μM dCTP-Cy5 and 190 μM dNTPs. The PCR reaction was carried out in a thermal cycler with the following procedure: 98℃ for 30 s, then 45 cycles of 98℃ for 10 s, 58℃ for 30 s and 72℃ for 1.5 min, then lastly 72℃ for 5 min. The product was purified using a NucleoSpin® PCR clean-up kit and the concentration determined by absorbance on a Nanodrop 2000c Spectrophotometer. DNA bands were visualized on a 1% agarose gel with ethidium bromide staining.

Revised version: For the dsDNA spring, a 5’ biotinylated forward primer (5’-/5Biosg/TGC CTC CGT GTA AGG GGG AT-3’) and a reverse primer with a 5’ overhang (5’-/GGG CCA TCG CCA ATT GGA GTA /idSp/ GTG AGT TAA AGT TGT ACT CGA GTT TGT GTC CAA GAA-3’) were used to create a 3009 bp dsDNA by PCR from plasmid Mus Musculus BicD2-sf-GFP (aa 25-425) in pet28a. The abasic Int 1’,2’-Dideoxyribose spacer (idSp) creates an overhang by terminating the polymerase. All oligonucleotides were purchased from IDT. Each 50 μL PCR reaction contained: 1x Phusion HF buffer, 198 μM dNTPs, 2 μM dCTP-Cy5, 0.5 μM primers, 3 ng template DNA and 1 U/50 μL HF Phusion Polymerase. Fluorescent dsDNA used in Figure 1 – figure supplement 4 had 10 μM dCTP-Cy5 and 190 μM dNTPs. The PCR reaction was carried out in a thermal cycler with the following procedure: 98℃ for 30 s, then 45 cycles of 98℃ for 10 s, 58℃ for 30 s and 72℃ for 1.5 min, then lastly 72℃ for 5 min. The product was purified using a NucleoSpin® PCR clean-up kit and the concentration determined by absorbance on a Nanodrop 2000c Spectrophotometer. DNA bands were visualized on a 1% agarose gel with ethidium bromide staining.

Motor-Microtubule-Tensiometer Assembly

Earlier version: Motors were bacterially expressed, purified and linked through its SNAP tag to an oligonucleotide (5’-/TAC TCC AAT TGG CGA TGG CCC / 3AmMC6T/-3’) complementary to the dsDNA overhang. Details of motor expression, purification and labeling, as well as tubulin biotinylation and polymerization are given in Supplementary Information. The DNA tensiometer was assembled on the microtubule as follows. The following three buffers are made on the same day of the experiment: C2AT (BRB80, 10 μM Taxol, 2 mM MgATP, 2 mg/mL Casein), 2AT (BRB80, 10 μM Taxol, 2 mM MgATP) and Imaging Solution (BRB80, 10 μM Taxol, 2 mg/mL Casein, 2 mM MgATP, 20 mM D-glucose, 0.02 mg/mL Glucose oxidase, 0.008 mg/mL Catalase, 0.5% BME and 2 mg/mL BSA). Full-length rigor kinesin was used to attach microtubules to the coverglass 65. Tensiometers were created in the flow cell using the following work flow: C2AT, 5 min > Rigor kinesin, 5 min> C2AT wash > BioMT, 5 min > 2AT > 8 nM Neutravidin, 5 min> 2AT > 10 nM Bio-dsDNA-Overhang, 5 min> C2AT > 4 nM KinesinMotor + 40nM Qdot-GBP (pre incubated in tube on ice for >15 min) in imaging solution, 10 min > Imaging solution wash. Note that because casein can contain free biotin, casein-free 2AT buffer was used during avidin-biotin binding steps. Following assembly, the Qdot connected to the motor was imaged on a custom-built TIRF microscope, described previously 84. Raw data were typically collected at 25 fps (range of 20-40 fps) on a Photometrics Prime 95B camera.

Revised version: Motors were bacterially expressed, purified and linked through its SNAP tag to an oligonucleotide (5’-/TAC TCC AAT TGG CGA TGG CCC / 3AmMC6T/-3’) complementary to the dsDNA overhang. Details of motor expression, purification and labeling, as well as tubulin biotinylation and polymerization are given in Supplementary Information. The DNA tensiometer was assembled on the microtubule as follows. The following three buffers are made on the same day of the experiment: C2AT (BRB80, 10 μM Taxol, 2 mM MgATP, 2 mg/mL Casein), 2AT (BRB80, 10 μM Taxol, 2 mM MgATP) and Imaging Solution (BRB80, 10 μM Taxol, 2 mg/mL Casein, 2 mM MgATP, 20 mM D-glucose, 0.02 mg/mL Glucose oxidase, 0.008 mg/mL Catalase, 0.5% BME and 2 mg/mL BSA). Full-length rigor kinesin was used to attach microtubules to the coverglass 77.

Added in the revised version: Tensiometers were created in the flow cell using the following work flow: C2AT, 5 min > Rigor kinesin, 5 min> C2AT wash > BioMT, 5 min > 2AT > 8 nM Neutravidin, 5 min> 2AT > 10 nM Bio-dsDNA-Overhang, 5 min> C2AT > 4 nM KinesinMotor + 40nM Qdot-GBP (pre incubated in tube on ice for >15 min) in imaging solution, 10 min > Imaging solution wash. Note that because casein can contain free biotin, casein-free 2AT buffer was used during avidin-biotin binding steps. Following assembly, the Qdot connected to the motor was imaged on a custom-built TIRF microscope, described previously 94. Raw data were typically collected at 25 fps (range of 20-40 fps) on a Photometrics Prime 95B camera.

Data Analysis

Earlier version: Movies were uploaded into FIESTA software 85 and Qdot intensities were tracked using a symmetric 2-D gaussian function to obtain x,y,t data for each Qdot. When drift correction was needed, TetraSpeck™ Fluorescent Microspheres (Thermo) and immobile Qdots were used as fiducial markers. The smallest position errors at stall in FIESTA fitting were 3-4 nm, which matched the positional error of Qdot-labeled motors stuck to microtubules in AMPPNP. Points with position errors greater than 20 nm were excluded because they often involved clearly spurrious position estimates. Notably, many tensiometers had small segments of missing data due to the Qdot fluctuating out of the TIRF field or blinking; these occurred most often during periods when the motors were detached from the microtubule.

Revised version: Movies were uploaded into FIESTA software 95 and Qdot intensities were tracked using a symmetric 2-D gaussian function to obtain x,y,t data for each Qdot. When drift correction was needed, TetraSpeck™ Fluorescent Microspheres (Thermo) and immobile Qdots were used as fiducial markers. The smallest position errors at stall in FIESTA fitting were 3-4 nm, which matched the positional error of Qdot-labeled motors stuck to microtubules in AMPPNP. Points with position errors greater than 20 nm were excluded because they often involved clearly spurious position estimates. Notably, many tensiometers had small segments of missing data due to the Qdot fluctuating out of the TIRF field or blinking; these occurred most often during periods when the motors were detached from the microtubule.

After obtaining X and Y positions of linear motor tracks in Fiesta, we rotated and translated the data in Excel to generate X versus t traces. The apparent origin was determined by averaging the points where the motor is fluctuating on its tether. In rare instances where no fluctuation was observed, the approximate origin was calculated by averaging the starting positions of all the ramps within the tensiometer (Figure 1). We then measured the ramp time, distance traveled, stall durations, reattachment times and starting positions. Tensiometers occasionally ended with the Qdot signal going dark, denoting either bleaching or failure of the Qdot-motor or motor-DNA connection. Notably, no clear instances of motor-Qdots walking past the plateau point (denoting the tensiometer breaking) were observed. Stalls that terminated due to the tensiometer going dark or the video ending were excluded from analysis.

Stochastic Modeling

Kinesin run and stall durations were simulated by using a modified version of published stochastic model of kinesin stepping 23,30. A motor is either in a strongly-bound state or a weakly-bound state (Figure 5A). At each timepoint, a motor in the strongly bound state can transition into the weakly-bound state with a first order transition rate constant, kS-W or step backward by 8 nm with a constant rate, kback = 3 s-1. A motor in the weakly-bound state can complete an 8-nm forward step by transitioning back to the strongly-bound state with rate constant, kW-S, or it can disengage from the microtubule with transition rate, kslip.

Earlier version: For simplicity, we set kS-W = kW-S at zero load. Load-dependent transition rates were defined as: where k0 is the unloaded transition rate, δ is the characteristic distance parameter and kBT is the Boltzmann’s constant multiplied by the absolute temperature, equal to 4.1 pN-nm at 25° C. Stall force was set to 6 pN. Unloaded and stall durations were simulated by starting the motor in strongly-bound state and continuing until it transitioned into the slip state, with 1000 simulations for each condition. Restart times were simulated by starting the motor in the slip state. From there, the motor can reengage by transitioning into the strongly-bound state or switch to one of two detached states having either a slow or fast recovery rate. Restart simulations were run for 10,000 iterations. Model parameters were constrained by experimental data, described fully in Supplementary Methods.

Revised version: For simplicity, we set kS-W = kW-S at zero load. Load-dependent transition rates were defined as: k(F)=k0∗eFδkBT where k0 is the unloaded transition rate, δ is the characteristic distance parameter and kBT is the Boltzmann’s constant multiplied by the absolute temperature, equal to 4.1 pN-nm at 25° C. Stall force was set to 6 pN. Unloaded and stall durations were simulated by starting the motor in strongly-bound state and continuing until it transitioned into the slip state, with 1000 simulations for each condition. Restart times were simulated by starting the motor in the slip state. From there, the motor can reengage by transitioning into the strongly-bound state or switch to one of two detached states having either a slow or fast recovery rate. Restart simulations were run for 10,000 iterations. Model parameters were constrained by experimental data, described fully in Supplementary Methods.